Before you turn this problem in, make sure everything runs as expected. First, **Restart and Run All** (in the menubar in colab, select Runtime$\rightarrow$Restart and Run all) 

Make sure you only fill in any place that says `YOUR CODE HERE` and do not make any other changes to the code. If you add any code cells for your own testing. Please delete them before submission.

---

```python
import numpy as np
from tensorflow import keras
from keras.datasets import mnist
```

Output:
```
Using TensorFlow backend.
```

## 1 
### Load Data and Rescale
Write a function that loads data, scales the image values by 1/255 and returns train and test sets 

<!-- grade_id: cell-8e7945a2fa914c65, points: 0, grade: false, solution: true -->
```python
def load():
  """
  Inputs:
    None
  
  Outputs:
    x_train : numpy array, training data shape (60000, 28, 28)
    x_test : numpy array, test data shape (10000, 28, 28)
    y_train : numpy array, training labels shape (60000, )
    y_test : numpy array, training labels shape (10000, )
  """
  
  # YOUR CODE HERE
  from keras.utils import to_categorical
  from keras.preprocessing.image import ImageDataGenerator
  (trainX, trainY), (testX, testY) = mnist.load_data()

  width, height, channels = trainX.shape[1], trainX.shape[2], 1
  trainX = trainX.reshape((trainX.shape[0], width, height, channels))
  testX = testX.reshape((testX.shape[0], width, height, channels))

  trainY = to_categorical(trainY)
  testY = to_categorical(testY)

  datagen = ImageDataGenerator(rescale=1.0/255.0)
 
  train_iterator = datagen.flow(trainX, trainY, batch_size=64)
  test_iterator = datagen.flow(testX, testY, batch_size=64)
  for i,j in train_iterator:
    x_train,y_train = i,j
  for a,b in test_iterator:
    x_test,y_test = a,b
  return x_train, x_test, y_train, y_test

```

<!-- grade_id: cell-a887bf40f92702ea, points: 5, grade: true, solution: false -->
```python
data = load()
assert data[0].shape == (60000, 28, 28)
print('Test passed', '\U0001F44D')

```

## 2
Different numbers consume different amounts of ink to write. So the total intensity of pixels is a feature which can separate different numbers. 

Write a function to calculate total intensity of pixels

<!-- grade_id: cell-529e6ce8f98a353f, points: 0, grade: false, solution: true -->
```python
def total_intensity(im):
  """
  Inputs:
  im : numpy array of shape (28, 28), image

  Outputs:
  tot_intensity : int, total intensity of pixels (sum of values of all pixels)
  """
  
  # YOUR CODE HERE
  
  
  return tot_intensity
```

<!-- grade_id: cell-776b3468dd622086, points: 5, grade: true, solution: false -->
```python
assert total_intensity(np.zeros((28, 28))) == 0
print('Test passed', '\U0001F44D')

```

## 3
Different numbers have different amounts of symmetry about Y axis. If we flip image of a number along the Y axis and subtract it from the original, the image that we see gives us an indication of the level of antisymmetry. Total intensity of this image can be a measure of anitsymmetry.

Find antisymmetry about X and Y axes.

<!-- grade_id: cell-043d67b64b813df1, points: 0, grade: false, solution: true -->
```python
def antisymmetries(im):
  """
  Inputs:
    im : numpy array of shape (28, 28), image

  Outputs:
    x_as : int, antisymmetry measure along x-axis (along width)
    y_as : int, antisymmetry measure along y-axis (along height)
  """
  
  # YOUR CODE HERE
  
  
  return x_as, y_as
```

<!-- grade_id: cell-eabc516d174a33aa, points: 5, grade: true, solution: false -->
```python
assert antisymmetries(np.ones((28, 28))) == (0, 0)
print('Test passed', '\U0001F44D')

```

## 4 
Create matrix of features for each image. Consider total intensity, and antisymmetries about x and y as the features. (Also take into account the bias term)

<!-- grade_id: cell-e6c58ace935183ba, points: 0, grade: false, solution: true -->
```python
def createA(x):
  """
  Inputs:
    x : numpy array of shape (N, 28, 28) , set of images
  
  Outputs:
    A : numpy array of shape (N, 4), array of features corresponding to each image
  """
  
  # YOUR CODE HERE
  
  
  return A
```

<!-- grade_id: cell-267844b01342c9d3, points: 5, grade: true, solution: false -->
```python
x_train = load()[0]
A = createA(x_train[: 5])
assert A.shape == (5, 4)
print('Test passed', '\U0001F44D')

```

## 5
Train a linear model and find accuracy 
- Create A matrix
- One-hot encode labels
- Train using pseudo-inverse and min-norm solutions
- Predict using trained model and find accuracy


<!-- grade_id: cell-fcf80d17d94b151f, points: 0, grade: false, solution: true -->
```python
def train_w(x_train, y_train):
  """
  Inputs:
    x_train : Numpy array of images (N, 28, 28)
    y_train : Numpy array of labels (N,)
  Outputs:
    w : numpy array (4, )
    train_accuracy : float (between 0 and 1)
  """
  
  # YOUR CODE HERE
  
  
  return w, train_accuracy
```

<!-- grade_id: cell-48d2d1cdd5aa22d1, points: 5, grade: true, solution: false -->
```python
x_train, x_test, y_train, y_test = load()
fives_and_zeros = (np.logical_or(y_train == 5, y_train == 0))
x_train_fives_and_zeros = x_train[fives_and_zeros]
y_train_fives_and_zeros = y_train[fives_and_zeros]
w, train_accuracy = train_w(x_train_fives_and_zeros, y_train_fives_and_zeros)

assert train_accuracy > 0.70  ## Better than random guess
print('Test passed', '\U0001F44D')

```