Before you turn this problem in, make sure everything runs as expected. First, **Restart and Run All** (in the menubar in colab, select Runtime$\rightarrow$Restart and Run all) 

Make sure you only fill in any place that says `YOUR CODE HERE` and do not make any other changes to the code. If you add any code cells for your own testing. Please delete them before submission.

---

- 5 marks for each question

# Q1

### Reverse words in string
Write a function that takes a string of words and returns a string with the words placed in the reverse order.

<!-- grade_id: cell-20abd36167cf1735, points: 0, grade: false, solution: true -->
```python
def reverse_words(s):
    """
    Inputs:
        s: string, string of words
    
    Outputs:
        s_rev : string, string of words reversed
    """
    s_rev = None
    
    # YOUR CODE HERE
    a = s.split(' ')
    a.reverse()
    s_rev = ' '.join(a)
    
    
    
    return s_rev
```

<!-- grade_id: cell-ea57ec0efd78ac56, points: 5, grade: true, solution: false -->
```python
assert reverse_words("Hello World") == "World Hello"
print('Test passed', '\U0001F44D')

```

Output:
```
Test passed 👍
```

# Q2

### Equal Sides Of An Array
Find index of list where sum of all elements to the left of the index and to the right of the index are equal. Return -1 otherwise.  <br>
Eg., ```arr = [1, 2, 3, 1, 1, 1]``` For index ```2```, left-sum is ```1+2+3 = 6``` and right sum is ```3+1+1+1=6```. So function should return ```2```. If there is no such index, function should return ```-1```

<!-- grade_id: cell-ef233da2ae2c5df8, points: 0, grade: false, solution: true -->
```python
def equal_sides(arr):
    """
    Inputs:
        arr: list of ints
    Outputs:
        ind: int, index of list where left-sum is equal to right-sum, -1 otherwise
    """
    ind = None
    
    # YOUR CODE HERE
    lst = []
    for i in arr:
        if sum(arr[:i+1]) == sum(arr[i:]):
            ind = i
            lst.append(i)
    if len(lst) == 0:
        ind = -1
    
    
    return ind
        
```

<!-- grade_id: cell-18e8be9f20b65b14, points: 5, grade: true, solution: false -->
```python
assert equal_sides([1, 2, 3, 1, 1, 1]) == 2
print('Test passed', '\U0001F44D')
```

Output:
```
Test passed 👍
```

# Q3

### K - palindrome
A string is a k-palindrome, if it becomes a palindrome on removing atmost k characters from the string.
For eg., ```'123431'``` is a palindrome on removing 1 character ```2``` from the string. So it is a 1-palindrome. 

<!-- grade_id: cell-06c0da9eb77e7de3, points: 0, grade: false, solution: true -->
```python
def k_palindrome(s, k):
    """
    Inputs:
        s: string
        k: int
    Outputs:
        is_k_palindrome: bool, True if s is k-palindrome, False otherwise
    """
    
    # YOUR CODE HERE
    s_new = s[::-1]

    n = len(s)

    
    a = [[0 for x in range(n + 1)] for y in range((n + 1))]

    
    for i in range(n + 1):
        for j in range(n + 1):
            
            if i == 0 or j == 0:
                a[i][j] = i + j

            
            elif s[i - 1] == s_new[j - 1]:
                a[i][j] = a[i - 1][j - 1]

            
            else:
                a[i][j] = 1 + min(a[i - 1][j], a[i][j - 1])

    return a[n][n] <= 2 * k
    
    
```

<!-- grade_id: cell-87e4402b5fa2e852, points: 5, grade: true, solution: false -->
```python
assert k_palindrome('123431', 1) == True
print('Test passed', '\U0001F44D')
```

Output:
```
Test passed 👍
```

# Q4

### N-digit sum
Given ```n``` and ```s```, find all n-digit numbers whose sum of digits adds upto ```s```.
Eg., for ```n = 3``` and ```s = 26```, the possibilities are ```998, 989, 899``` <br>
Return all the numbers as a sorted list. Return an empty list if no such numbers exist

<!-- grade_id: cell-b767907cfc0654e5, points: 0, grade: false, solution: true -->
```python
def n_digit_sum(n, s):
    """
    Inputs:
        n : int, number of digits
        s : int, sum of digits
    Outputs:
        arr : list, sorted list of n-digit numbers where sum of digits of each number is s
    """
    arr = None
    
    # YOUR CODE HERE
    min_ = int('1' + '0'*(n-1))
    max_ = int('9' * n)
    arr = []
    for i in range(min_,(max_+1)):
        x = [int(j) for j in list(str(i))]
        if sum(x) == s:
            arr.append(i)
    
    return arr
    
```

<!-- grade_id: cell-f1044d16ec02246e, points: 5, grade: true, solution: false -->
```python
assert n_digit_sum(3, 27) == [999]
print('Test passed', '\U0001F44D')
```

Output:
```
Test passed 👍
```

# Q5

### Keep
Write a function ```keep``` that takes two inputs - 
- a list of elements ```l```
- a function f which returns a ```bool``` (either ```True``` or ```False```) for each element of ```l```

You need to return ```keep_list``` which contains all the elements of l for which ```f``` returns ```True```

<!-- grade_id: cell-3c66934816e806d3, points: 0, grade: false, solution: true -->
```python
def keep(l, f):
    """
    Inputs:
        l: list
        f: function, returns bool for each element of l
    Outputs:
        keep_list: list
    """
    keep_list = None
    
    # YOUR CODE HERE
    keep_list = []
    for i in range(len(l)):
        if f(l[i]) == True:
            keep_list.append(l[i])
    
    
    return keep_list
```

<!-- grade_id: cell-917e6637446d38c8, points: 5, grade: true, solution: false -->
```python
def is_even(x):
    return x%2 == 0
assert keep([1, 2, 3, 4, 5, 6, 8, 9], is_even) == [2, 4, 6, 8]
print('Test passed', '\U0001F44D')
```

Output:
```
Test passed 👍
```