Before you turn this problem in, make sure everything runs as expected. First, **Restart and Run All** (in the menubar in colab, select Runtime$\rightarrow$Restart and Run all) 

Make sure you only fill in any place that says `YOUR CODE HERE` and do not make any other changes to the code. If you add any code cells for your own testing. Please delete them before submission.

---

```python
import numpy as np
```

## 1

### Micro Average Precision
Given y_true and y_pred, find the micro averaged precision score.

micro_precision = $\frac{TP_{\text{all}}}{TP_{\text{all}} + FP_{\text{all}}} $

where <br>
$TP_{\text{all}}$ is the total number of true positive predictions considering all classes and <br>
$FP_{\text{all}}$ is the total number of false positive predictions considering all classes

<!-- grade_id: cell-5a15a882cc098e02, points: 0, grade: false, solution: true -->
```python
def micro_avg_precision(y_true, y_pred):
  """
  Inputs:
    y_true: 1-D numpy array, true labels
    y_pred: 1-D numpy array, predicted labels
  Outputs:
    micro_precision: float
  """
  
  # YOUR CODE HERE
  from sklearn.metrics import precision_score
  micro_precision = precision_score(y_true,y_pred,average='micro')
  
  return micro_precision
```

<!-- grade_id: cell-b75352c8cdb22c1f, points: 5, grade: true, solution: false -->
```python
y_true = np.array(['cat', 'dog', 'cat', 'human'])
y_pred = np.array(['cat', 'human', 'dog', 'cat'])
assert micro_avg_precision(y_true, y_pred) == 0.25
print('Test passed', '\U0001F44D')

```

Output:
```
Test passed 👍
```

## 2

### Macro Average Recall
Given y_true and y_pred, find the macro averaged recall score where macro averaged recall is the unweighted average of the recall score for each class.

<!-- grade_id: cell-8431a49161fa30ed, points: 0, grade: false, solution: true -->
```python
def macro_avg_recall(y_true, y_pred):
  """
  Inputs:
    y_true: 1-D numpy array, true labels
    y_pred: 1-D numpy array, predicted labels
  Outputs:
    macro_all: float
  """
  
  # YOUR CODE HERE
  from sklearn.metrics import recall_score
  macro_recall = recall_score(y_true,y_pred,average='macro')
  
  return macro_recall
```

<!-- grade_id: cell-fc0e57143e4c9af8, points: 5, grade: true, solution: false -->
```python
y_true = np.array(['cat', 'dog', 'cat', 'human'])
y_pred = np.array(['cat', 'human', 'dog', 'cat'])
assert np.round(macro_avg_recall(y_true, y_pred), 2) == 0.17
print('Test passed', '\U0001F44D')


```

Output:
```
Test passed 👍
```

## 3

### Weighted F1 Score
Given y_true and y_pred, find the wieghted f1 score where the weighted f1 score is the weighted average of the f1 score of each class weighted by number of true samples of that class.

<!-- grade_id: cell-f4359c7dfa348934, points: 0, grade: false, solution: true -->
```python
def weighted_avg_f1(y_true, y_pred):
  """
  Inputs:
    y_true: 1-D numpy array, true labels
    y_pred: 1-D numpy array, predicted labels
  Outputs:
    weighted_f1: float
  """
  
  # YOUR CODE HERE
  from sklearn.metrics import f1_score
  weighted_f1 = f1_score(y_true,y_pred,average='weighted')
  
  return weighted_f1
```

<!-- grade_id: cell-3590ecd4d88051f7, points: 5, grade: true, solution: false -->
```python
y_true = np.array(['cat', 'dog', 'cat', 'human'])
y_pred = np.array(['cat', 'human', 'dog', 'cat'])
assert weighted_avg_f1(y_true, y_pred) == 0.25
print('Test passed', '\U0001F44D')


```

Output:
```
Test passed 👍
```

## 4

### Multi Label Confusion Matrix
Given a multi-label ground truth and predictions, create a multi-label confusion matrix.

<!-- grade_id: cell-b550bf4f356e9e65, points: 0, grade: false, solution: true -->
```python
def multi_label_confusion_matrix(y_true, y_pred):
  """
  Inputs:
    y_true : numpy array of shape (N, C) where N is number of samples and C is number of classes. (C is multi-hot, 1 if label is present, 0 otherwise)
    y_pred : numpy array of shape (N, C) where N is number of samples and C is number of classes. (C is multi-hot, 1 if label is present, 0 otherwise)
  
  Outputs:
    ml_cm : numpy array of shape (C, 2, 2) where C is number of classes and each 2x2 matrix is a confusion matrix for that class where rows are true labels and columns are predicted labels. 
            Order of classes should be sorted by their index in y_true/y_pred
  """
  
  # YOUR CODE HERE
  from sklearn.metrics import multilabel_confusion_matrix
  ml_cm = multilabel_confusion_matrix(y_true, y_pred)
  
  return ml_cm

```

<!-- grade_id: cell-51413a727505af1c, points: 5, grade: true, solution: false -->
```python
y_true = np.array([[1, 0, 1],
                   [0, 1, 0]])
y_pred = y_pred = np.array([[1, 0, 0],
                 [0, 1, 1]])
assert np.all(multi_label_confusion_matrix(y_true, y_pred) == np.array([[[1, 0],[0, 1]],[[1, 0],[0, 1]],[[0, 1],[1, 0]]]))
print('Test passed', '\U0001F44D')


```

Output:
```
Test passed 👍
```

## 5

### Hamming Score
Hamming Score or Hamming Loss is the ratio of total number of incorrect label predictions and the total number of label predictions for a multi-label classification problem. (https://en.wikipedia.org/wiki/Multi-label_classification#Statistics_and_evaluation_metrics

<!-- grade_id: cell-7f827bcaac0f444b, points: 0, grade: false, solution: true -->
```python
def hamming_score(y_true, y_pred):
  """
  Inputs:
    y_true : numpy array of shape (N, C) where N is number of samples and C is number of classes. (C is multi-hot, 1 if label is present, 0 otherwise)
    y_pred : numpy array of shape (N, C) where N is number of samples and C is number of classes. (C is multi-hot, 1 if label is present, 0 otherwise)

  Outputs:
    h_loss : float, hamming score
  """
  
  # YOUR CODE HERE
  from sklearn.metrics import hamming_loss
  h_loss = hamming_loss(y_true,y_pred)
  
  return h_loss

```

<!-- grade_id: cell-63323cf2023fb2c9, points: 5, grade: true, solution: false -->
```python
y_true = np.array([[0, 1], [1, 1]])
y_pred = np.zeros((2, 2))
assert hamming_score(y_true, y_pred) == 0.75
print('Test passed', '\U0001F44D')

```

Output:
```
Test passed 👍
```

```python

```