Before you turn this problem in, make sure everything runs as expected. First, **Restart and Run All** (in the menubar in colab, select Runtime$\rightarrow$Restart and Run all) 

Make sure you only fill in any place that says `YOUR CODE HERE` and do not make any other changes to the code.<br> If you add any code cells for your own testing. Please delete them before submission.

---

```python
import math
```

- Five marks for each question.

## Q1

### Third Largest Element
Write a function that returns the third largest element of a list of integers. (Assume all elements in l will be distinct)

<!-- grade_id: cell-08e7132823c18a19, points: 0, grade: false, solution: true -->
```python
def third_largest_element(l):
    """
    Inputs:
        l: list of ints, length of list >=3
    
    Outputs:
        third_largest: int, third largest integer in l
    """
    
    # YOUR CODE HERE
    if len(l) >= 3:
      l.sort()
      third_largest = l[-3]
    
    
    return third_largest
```

<!-- grade_id: cell-67c5d87e34bb12bc, points: 5, grade: true, solution: false -->
```python
assert type(third_largest_element([1, 2, 3])) == type(4)

```

## Q2

### Taylor Series
Taylor Series expansion of $e^x$ upto $n$ terms is as follows <br>
$$
e^x = \sum_{i=0}^{n}{\frac{x^i}{i!}}
$$
where
$i!$ is factorial of $i$ i.e., $1*2*3*\ldots*i$

Write a function that takes arguments $n$ and $x$ and finds value of taylor expansion of $e^x$ upto $n$ terms

<!-- grade_id: cell-2b724d2d1a657c8f, points: 0, grade: false, solution: true -->
```python
def taylor_expansion_ex(x, n):
    """
    Inputs:
        n: int, n>=0
        x: float
    
    Outputs:
        taylor_exp : float, taylor_expansion of e^x upto n terms
    """
    
    # YOUR CODE HERE
    ans = 0
    for i in range(0,n):
        ans += x**i / math.factorial(i)
    taylor_exp = ans 
    
    return taylor_exp
```

<!-- grade_id: cell-601893a812144e25, points: 5, grade: true, solution: false -->
```python
assert type(taylor_expansion_ex(1, 1)) == type(2.)

```

## Q3

### Armstrong numbers
Armstrong number is a number that is equal to the sum of cubes of its digits. <br>
For eg., $153 = 1^3 + 5^3 + 3^3$

Write a function that takes an argument $n$ and returns list of armstrong numbers between $0$ and $n$ (including $0$ and $n$)(List should be sorted in ascending order)

<!-- grade_id: cell-4bd7d534db20e974, points: 0, grade: false, solution: true -->
```python
def armstrong_numbers(n):
    """
    Inputs:
        n: int
        
    Outputs:
        arm_nums: sorted list of ints, list of armstrong numbers between 0 to n 
    """
    
    # YOUR CODE HERE
    def arm_no(x):
      b = [int(i) for i in str(x)]
      c = [i**3 for i in b]
      return sum(c)
    lst = []
    for i in range(0,n+1):
      lst.append(arm_no(i))
    lst.sort()
    arm_nums = lst
    return arm_nums
```

<!-- grade_id: cell-f39a16c6615c36c1, points: 5, grade: true, solution: false -->
```python
assert type(armstrong_numbers(1)) == list

```

## Q4

### Remove all occurences
Given a list $l$ and element $e$, remove all occurences of $e$ from $l$ and return it. 

<!-- grade_id: cell-f61ee94d18289e65, points: 0, grade: false, solution: true -->
```python
def remove_occurences(l, e):
    """
    Inputs:
        l: list
        e: an element of list l (can be of type int, float or str)
    Outputs:
        l_clean: list, same as list l but with all occurences of e removed from it
    """
    
    # YOUR CODE HERE
    while e in l:
      l.remove(e)
    l_clean = l
    
    return l_clean
```

<!-- grade_id: cell-a8521f22df5c2000, points: 5, grade: true, solution: false -->
```python
assert type(remove_occurences([1, 2, 3, 4], 2)) == list

```

## Q5

## Adder function
Write a function that takes an input $a$ and returns a function $f$. <br>
$f$ should take an input $n$ and return the sum of $a$ and $n$. 


<!-- grade_id: cell-4509d678d995333b, points: 0, grade: false, solution: true -->
```python
def get_adder(a):
    """
    Inputs:
        a: int/float
    Outputs:
        a_adder: function, it is a function that takes an input n, adds a to it and returns it
    """
    
    # YOUR CODE HERE
    def f(n):
      return n + a
    a_adder = f
    return a_adder
```

<!-- grade_id: cell-17b02b44087f0987, points: 5, grade: true, solution: false -->
```python
assert type(get_adder(2)) == type(lambda x: x)

two_adder = get_adder(2)
assert two_adder(3) == 5

```